## Problem 2 - Temperature classifier (3 points)

In [Problem 2 of Exercise 3 last week](https://github.com/Geo-Python-2023/Exercise-3/blob/master/Exercise-3-problem-2.ipynb) we practiced how to classify temperatures into four different classes: cold, slippery, comfortable, and warm. Let's continue working with the same idea.

Here, you should create a function called `temp_classifier` that accepts a temperature value in Celsius that will be reclassified into integer numbers 0-3 based on following criteria:

| Return value | Classification criteria                                                  |
| :----------: | :----------------------------------------------------------------------- |
| 0            | Temperatures below -2 degrees Celsius                                    |
| 1            | Temperatures equal or warmer than -2, but less than +2 degrees Celsius   |
| 2            | Temperatures equal or warmer than +2, but less than +15 degrees Celsius  |
| 3            | Temperatures equal or warmer than +15 degrees Celsius                    |

### Tips for completing this problem

- Use **exactly** the same variable names as in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.
- **Please do not**:

    - **Change the file names**. Do all of your editing in the provided `Exercise-4-problem-2.ipynb` file (this file).
    - **Copy/paste cells in this notebook**. We use an automated grading system that will fail if there are copies of code cells.
    - **Change the existing cell types**. You can add cells, but changing the cell types for existing cells (from code to markdown, for example) will also cause the automated grader to fail.

### Scores for this problem

**Your score on this problem will be based on following criteria:**

- Creating a new function called `temp_classifier` **(0.5 points)**
- Defining one parameter called `temp_celsius` in your `temp_classifier` function **(0.5 points)**
- Reclassifying input temperatures based on the criteria in the above table and returning the reclassified value as an integer number (i.e., 0, 1, 2, or 3) **(1 point)**
- Adding comments in your code and a docstring that explains how to use your `temp_classifier` function (i.e., you should write the purpose of the function, parameters, and returned values) **(1 point)**
- Uploading your notebook to your GitHub repository for this week's exercise

## AI tool usage agreement

**Enter your name (and that of your partner) in the cell below** to confirm that you have followed the [course guidelines on the use of AI tools](https://geo-python-site.readthedocs.io/en/latest/course-info/ai-tools.html) and understand that misuse of AI tools is considered cheating.

<!-- grade_id: E4-P2-AI-LLM_statement, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Part 1 (2 points)

Create a function in the cell below to classify temperatures as directed.

- Your function should be called `temp_classifier`
- It should have a single parameter called `temp_celsius`
- It should properly classify temperatures for a given temperature input value

<!-- grade_id: problem_2_ag_p1_temp_classifier, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def temp_classifier(temp_celsius):
    if temp_celsius < -2:
      return 0
    elif -2 <= temp_celsius < 2:
      return 1
    elif 2 <= temp_celsius < 15:
      return 2
    else:
      return 3
```

Use your newly created function to print the answers to the following questions:

```python
# 1. What is the class value for 16.5 degrees Celsius?
print(f"the class value for 16.5 degrees is {temp_classifier(16.5)}")

# 2. What is the class value for +2 degrees Celsius?
print(f"the class value for +2 degrees is {temp_classifier(2)}")

```

Output:
```
the class value for 16.5 degrees is 3
the class value for +2 degrees is 2
```

### Problem 2 tests

Run the code cells below to check your answer. Note that you are not able to modify the test cells. Your solution is on the right track if you don't see any errors when running these code cells.

<!-- grade_id: problem_2_ag_p1_temp_classifier_test1, points: 0.5, grade: true, solution: false -->
```python
import inspect

# Check that function exists
assert inspect.isfunction(temp_classifier)
```

<!-- grade_id: problem_2_ag_p1_temp_classifier_test2, points: 0.5, grade: true, solution: false -->
```python
# Check that the function has a single parameter and the pamameter name is correct
params = list(inspect.signature(temp_classifier).parameters.keys())
assert len(params) == 1
assert params[0] == 'temp_celsius'
```

Check that the function produces correct answers for selected values:

<!-- grade_id: problem_2_ag_p1_temp_classifier_test3, points: 0.25, grade: true, solution: false -->
```python
# 1. What is the class value for 16.5 degrees (Celsius)?
assert temp_classifier(16.5) == 3, 'Wrong class'
print("ok :)")
```

Output:
```
ok :)
```

<!-- grade_id: problem_2_ag_p1_temp_classifier_test4, points: 0.25, grade: true, solution: false -->
```python
# 2. What is the class value for +2 degrees (Celsius)?
assert temp_classifier(2) == 2, 'Wrong class'
print("ok :)")
```

Output:
```
ok :)
```

<!-- grade_id: problem_2_ag_p1_temp_classifier_test5, points: 0.25, grade: true, solution: false -->
```python
# 3. What is the class value for +1 degrees (Celsius)?
assert temp_classifier(1) == 1, 'Wrong class'
print("ok :)")
```

Output:
```
ok :)
```

<!-- grade_id: problem_2_ag_p1_temp_classifier_test6, points: 0.25, grade: true, solution: false -->
```python
# 4. What is the class value for -5 degrees (Celsius)?
assert temp_classifier(-5) == 0, 'Wrong class'
print("ok :)")
```

Output:
```
ok :)
```

## Part 2 (1 point)

Check your code.

- Are all the variable names correct?
- Does your code contain enough comments?
- Does your function have an informative dosctring?

You can add any feedback you have about this problem in the cell below.

<!-- grade_id: problem_2_manual_check, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Part 3 (*optional*, 0 points) - **AI-LLM OK**

Using [ChatGPT](https://openai.com/chatgpt), create a new function in the cell below to classify temperatures as directed at the start of this problem.

In order to do this, we request that you post both the prompt you provide to ChatGPT and the code it produces in the cells below. As your prompts may need to be refined to get your code working, please report each prompt so we can see how you arrived at your final version of the function.

- Your function should be called `ai_temp_classifier`
- It should have a single parameter called `temp_celsius`
- It should properly classify temperatures for a given temperature input value

### Prompt 1

Enter the text of the first prompt you used for this part of the exercise below.

<!-- grade_id: problem_2_prompt1, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

Enter the code generated by your first prompt below.

<!-- grade_id: problem_2_prompt1_code, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Looking at the function produced by ChatGPT do you see any problems compared to your solution in Part 1? Please comment below about this.

<!-- grade_id: problem_2_prompt1_comments, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

Test out your code by running the tests below.

<!-- grade_id: problem_2_ag_p3_ai_temp_classifier_test1, points: 0, grade: true, solution: false -->
```python
# 1. What is the class value for 16.5 degrees (Celsius)?
assert ai_temp_classifier(16.5) == 3, 'Wrong class'
print("ok :)")
```

<!-- grade_id: problem_2_ag_p3_ai_temp_classifier_test2, points: 0, grade: true, solution: false -->
```python
# 2. What is the class value for +2 degrees (Celsius)?
assert ai_temp_classifier(2) == 2, 'Wrong class'
print("ok :)")
```

<!-- grade_id: problem_2_ag_p3_ai_temp_classifier_test3, points: 0, grade: true, solution: false -->
```python
# 3. What is the class value for +1 degrees (Celsius)?
assert ai_temp_classifier(1) == 1, 'Wrong class'
print("ok :)")
```

<!-- grade_id: problem_2_ag_p3_ai_temp_classifier_test4, points: 0, grade: true, solution: false -->
```python
# 4. What is the class value for -5 degrees (Celsius)?
assert ai_temp_classifier(-5) == 0, 'Wrong class'
print("ok :)")
```

### Additional prompts

In the event your code does not pass all of the tests above, please enter the prompt(s) you used to refine the function in the cell below. If you used more than one prompt, you can format them as a Markdown list using the format demonstrated below:

```
1. First prompt
2. Second prompt
3. ...
```

If no additional prompts were needed, you can skip this.

<!-- grade_id: problem_2_additional_prompts, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

Enter the final version of your `ai_temp_classifier` function produced by ChatGPT in the cell below.

<!-- grade_id: problem_2_additional_prompts_code, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Test out your revised code using the tests below.

<!-- grade_id: problem_2_ag_p3_ai_temp_classifier_test5, points: 0, grade: true, solution: false -->
```python
# 1. What is the class value for 16.5 degrees (Celsius)?
assert ai_temp_classifier(16.5) == 3, 'Wrong class'
print("ok :)")
```

<!-- grade_id: problem_2_ag_p3_ai_temp_classifier_test6, points: 0, grade: true, solution: false -->
```python
# 2. What is the class value for +2 degrees (Celsius)?
assert ai_temp_classifier(2) == 2, 'Wrong class'
print("ok :)")
```

<!-- grade_id: problem_2_ag_p3_ai_temp_classifier_test7, points: 0, grade: true, solution: false -->
```python
# 3. What is the class value for +1 degrees (Celsius)?
assert ai_temp_classifier(1) == 1, 'Wrong class'
print("ok :)")
```

<!-- grade_id: problem_2_ag_p3_ai_temp_classifier_test8, points: 0, grade: true, solution: false -->
```python
# 4. What is the class value for -5 degrees (Celsius)?
assert ai_temp_classifier(-5) == 0, 'Wrong class'
print("ok :)")
```

Does your revised function now pass all of the tests?
- If yes, please comment below about your experience using ChatGPT and the benefits and drawbacks you see.
- If no, what kinds of problems do you see with the revised function produced by ChatGPT compared to your solution in Part 1? Please comment below about this.

<!-- grade_id: problem_2_additional_prompts_comments, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### Final comments

If you have any other comments about your experience using ChatGPT in this problem, please enter those below.

<!-- grade_id: problem_2_chatgpt_comments, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

#### Done!

That's it! Now you are ready to continue with [Problem 3](Exercise-4-problem-3.ipynb).