## Problem 1 - Simple temperature calculator (3 points)

In the first problem your aim is to create a function that converts the input temperature from degrees Fahrenheit to degrees Celsius. The conversion formula from Fahrenheit to Celsius can be found below.

\begin{equation}
  \Large
  T_{\mathrm{Celsius}} = \left( T_{\mathrm{Fahrenheit}} - 32 \right) / 1.8
\end{equation}

If everything in your script is working properly the following test case should work:

```python
>>> print(f"32 degrees Fahrenheit is equal to {fahr_to_celsius(32)} degrees Celsius.")
32 degrees Fahrenheit is equal to 0.0 degrees Celsius.
```

For this problem you will need to modify the code blocks below and add code to produce the desired outcomes. 

### Tips for completing this problem

- Use **exactly** the same variable names as in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.
- **Please do not**:

    - **Change the file names**. Do all of your editing in the provided `Exercise-4-problem-1.ipynb` file (this file).
    - **Copy/paste cells in this notebook**. We use an automated grading system that will fail if there are copies of code cells.
    - **Change the existing cell types**. You can add cells, but changing the cell types for existing cells (from code to markdown, for example) will also cause the automated grader to fail.

### Scores for this problem

**Your score on this problem will be based on following criteria:**

- Creating a function called `fahr_to_celsius`
- Defining the function to have one input parameter called `temp_fahrenheit`
- Creating a variable called `converted_temp` inside the function to which you should assign the conversion result (i.e., the input Fahrenheit temperature converted to Celsius)
- Returning the converted value from the function back to the user
- Answering some questions about functions at the end of this problem
- Adding comments in your code and a docstring that explains how to use your `fahr_to_celsius` function (i.e., you should write the purpose of the function, parameters, and returned values)
- Uploading your notebook to your GitHub repository for this week's exercise

## AI tool usage agreement

**Enter your name (and that of your partner) in the cell below** to confirm that you have followed the [course guidelines on the use of AI tools](https://geo-python-site.readthedocs.io/en/latest/course-info/ai-tools.html) and understand that misuse of AI tools is considered cheating.

<!-- grade_id: E4-P1-AI-LLM_statement, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Part 1 (2.5 points)

Create a function in the cell below to convert temperatures from Fahrenheit to Celsius.

- The function should be called `fahr_to_celsius`
- It should have one parameter called `temp_fahrenheit`
- It should return a single value, the converted temperature
- You should include some inline comments and a docstring

<!-- grade_id: problem_1_ag_p1_fahr_to_celsius, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Use your newly created function to print answers to the following questions:

```python
# 1. What is 48째 Fahrenheit in Celsius? 

# 2. What about 71째 Fahrenheit in Celsius?

```

### Problem 1 tests

Run the code cells below to check your answer. Note that you are not able to modify the test cells. Your solution is on the right track if you don't see any errors when running these code cells. If you receive an error, you need to change something in the code cells above.

<!-- grade_id: problem_1_ag_p1_fahr_to_celsius_test1, points: 1, grade: true, solution: false -->
```python
import inspect

# Check that function exists
assert inspect.isfunction(fahr_to_celsius), 'Fahr_to_celsius should be a function.'
```

<!-- grade_id: problem_1_ag_p1_fahr_to_celsius_test2, points: 0.5, grade: true, solution: false -->
```python
# Check that the function has a single parameter and the parameter name is correct
params = list(inspect.signature(fahr_to_celsius).parameters.keys())
assert len(params) == 1, 'The function should have one parameter'
assert params[0] == 'temp_fahrenheit', 'The parameter name should be "temp_fahrenheit".'
```

<!-- grade_id: problem_1_ag_p1_fahr_to_celsius_test3, points: 0.5, grade: true, solution: false -->
```python
# Check that the function produces correct answers for:
# 1. What is 48째 Fahrenheit in Celsius? 
assert round(fahr_to_celsius(48), 2) == 8.89
```

<!-- grade_id: problem_1_ag_p1_fahr_to_celsius_test4, points: 0.5, grade: true, solution: false -->
```python
# 2. What about 71째 Fahrenheit in Celsius?
assert round(fahr_to_celsius(71), 2) == 21.67
```

### Check your code

- Make sure you used the given variable names
- Check that you have added necessary comments to your code
- Check that your function has a docstring that describes what it does

## Part 2 (0.5 points)

We would like you to think about and answer the following questions based on the materials and ideas that you learned during the lecture.

1. Is the concept of function clear to you? If not, what do you not understand?
2. What are some of the benefits of using functions?
  
Write your answers below (**activate the cell by double clicking**).

<!-- grade_id: problem_1_mg_p2_answers, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

#### Done!

That's it! Now you are ready to continue with [Problem 2](Exercise-4-problem-2.ipynb).